---
title: "Lesson 4: Dataframes" 
author: "Copyright 2024 Psychology Department. Hebrew University of Jerusalem. All rights reserved"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: false


description: >
 <div style='direction: rtl;'>
 בשיעור זה נלמד לעבוד עם טבלאות בR. נלמד על חיתוכים, יצירת עמודות חדשות ועריכת עמודות קיימות. נלמד גם כיצד לייבא קבצים מסוג csv ולערוך אותם בR.
 </div>
       
runtime: shiny_prerendered  
editor_options: 
  markdown: 
    wrap: 72
---


```{=html}
<style>
h1, h2, h3, h4, h5, h6 {
  direction: rtl;
}
p {
  direction: rtl;
}

.text-block1 {
  direction: rtl;       /* Set text direction to right-to-left */
  text-align: right;
  background-color: #e7f3fe; /* Light grey background */
  padding: 10px;
  border-radius: 5px;
  border: 1px solid #ddd; /* Light border */
  margin: 10px 0; /* Space around the block */
}
  

</style>


```

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(Rcourse)

tutorial_options(
  exercise.timelimit = 60
)

knitr::opts_chunk$set(error = TRUE)
gradethis::gradethis_setup()
```

```{r prepare-df1}
df1 = data.frame(subject = 1:6,
                city = c("Jerusalem","Haifa","Tel-Aviv","Eilat","Jerusalem","Tel-Aviv"),
                self_efficacy_score = c(2.3, 2.5, 2.1, 3, 2.2, 2.6))
```

## טבלאות - היכרות

בשיעור הנוכחי נלמד לעבוד עם סוג נוסף של אובייקטים שנפוצים מאוד בעבודה
שלנו בR: טבלאות, או באנגלית - Dataframes

הטבלאות מורכבות משורות ועמודות, כמו באקסל, עם ערכים בתוך כל אחת מהמשבצות
של הטבלה. לכל עמודה יש כותרת, המתארת את המידע באותה עמודה. שורות שונות
בדרך כלל יבטאו תצפיות שונות, למשל - נתונים של נבדקים שונים או סבבים
שונים של מטלה. לפניכם דוגמה של טבלה המתארת את
הנתונים הדמוגרפיים והציונים במדד תחושת מסוגלות עצמית עבור שישה נבדקים
פיקטיביים.

```{r}
df1 = data.frame(subject = 1:6,
                city = c("Jerusalem","Haifa","Tel-Aviv","Eilat","Jerusalem","Tel-Aviv"),
                self_efficacy_score = c(2.3, 2.5, 2.1, 3, 2.2, 2.6))

print(df1)
```

כפי שניתן לראות בשורות הקוד המייצרות את הטבלה הנ"ל, טבלאות (Dataframes)
הן אסופה של מספר וקטורים בעלי אותו האורך - כשלכל וקטור מוצמדת כותרת.
במקרה הנ"ל הטבלה מורכבת מ3 וקטורים, כל אחד באורך 6 ערכים. לוקטור הראשון בחרנו את השם "subject" והוא מתאר את המספר הסידורי של כל נבדק, לשני קראנו "city" ובו מתועדת עיר המגורים של הנבדק ולוקטור השלישי הוספנו את הכותרת "self_efficacy_score" , שכן הוא כולל את ציוני הנבדקים במדד תחושת מסוגלות עצמית. בהתאם לכך
שמדובר בוקטורים - כל עמודה חייבת להכיל סוג אחד של נתונים בלבד- numeric,
character, boolean וכו', אך עמודות שונות באותה טבלה יכולות להכיל סוגים
שונים של מידע.


מעיון בטבלה אפשר לראות כי נבדק 1 גר בירושלים והוא בעל ציון 2.3 במדד תחושת המסוגלות
העצמית. לעומת זאת נבדק 2 גר בחיפה והוא בעל ציון של 2.5 במדד המסוגלות
העצמית.

```{r question1_5, echo = FALSE}

question_text(
  "מהו ציון המסוגלות העצמית של הנבדק הגר באילת?",
  answer("3", correct = TRUE),
  answer("3.0", correct = TRUE),
  allow_retry = TRUE,
    
  correct="נכון מאוד", 
  incorrect = "נסו שוב"
)

```

### יצירת טבלאות

הדרך הכי פשוטה לייצר טבלה היא באמצעות הפקודה `data.frame`, שמאפשרת לנו
לצרף יחד מספר וקטורים לכדי טבלה, כפי שעשינו בדוגמה הקודמת. זו לרוב לא תהיה הדרך בה נייצר את
הטבלאות שאיתן נעבוד, אבל חשוב להכיר גם אותה.

ניתן לייצר את מסגרת הנתונים באמצעות וקטורים שאנחנו מגדירים בתוך
הפונקציה או להשתמש בוקטורים שהוגדרו
מראש. שימו לב: העמודות חייבות להיות כולן באותו האורך. במידה ונגדיר עמודה חדשה באמצעות וקטור קצר או ארוך יותר,
נקבל הודעת שגיאה

::: text-block1
המקרה החריג, בו לא נקבל הודעת שגיאה הוא כשנתשמש בוקטור בעל אורך שהוא
בדיוק חצי, שליש, רבע וכו' מאורך שאר העמודות בטבלה. במקרים כאלו התכנה תשכפל את הוקטור
לפי מספר הפעמים הנדרש כדי להגיע לאורך התואם את מספר השורות הנדרש.
:::

שורות הקוד הבאות מייצרות טבלה זהה לטבלה הקודמת, אבל מגדירה את כל אחד מהוקטורים בנפרד ורק אח"כ מצרפת אותם לכדי טבלה.


`תרגיל` ערכו את הפונקציה המייצרת את הטבלה כך שתתווסף עמודה רביעית בשם `gender` ומלאו אותה בערכים
לבחירתכם.


```{r exercise1_5, exercise=TRUE, exercise.eval = FALSE }

subjects = 1:6
city_per_subject = c("Jerusalem","Haifa","Tel-Aviv","Eilat","Jerusalem","Tel-Aviv")
self_efficacy_per_subject =  c(2.3, 2.5, 2.1, 3, 2.2, 2.6)

df1 = data.frame(subject = subjects,
                 city =  city_per_subject,
                 self_efficacy_score = self_efficacy_per_subject)

print(df1)
```

```{r exercise1_5-check}
grade_this(
   if (! "gender" %in% names(user_object_get("df1"))){
     fail("לא נמצאה עמודה בשם gender")}
   else {pass("מעולה!")}
   )
```

## חיתוכים

### חיתוך באמצעות סוגריים מרובעים [ , ]



#### חיתוך משבצות בודדות



חיתוך ושליפה של ערכים מתוך טבלה נעשית באופן דומה מאוד לשליפת ערכים מתוך
וקטור - אבל כיוון שלטבלאות יש 2 מימדים (שורות ועמודות) נצטרך להשתמש ב2
אינדקסים כדי להתייחס למשבצת / משבצות שמעניינות אותנו. האינדקס הראשון (בצד השמאלי של הסוגריים) יתייחס לשורות והשני (בצד הימני) יתייחס לעמודות, כשבין שני האינדקסים מפריד
פסיק. למשל: החיתוך `df1[4,3]` מתייחס לשורה 4, עמודה 3 של הטבלה שנקראת `df1`.

בשורות הקוד הבאות אנחנו שולפים את הערך של העיר (עמודה 2) בה גר הנבדק
שבשורה השלישית ומדפיסים אותו.


`תרגיל:` ערכו את הקוד כך שישלוף וידפיס את ציון תחושת המסוגלות של הנבדק שבשורה ה5.

```{r exercise2_5, exercise=TRUE, exercise.eval = FALSE,  exercise.setup = "prepare-df1"}
print(df1)

x = df1[3,2]
print(x)
```

```{r exercise2_5-check}
grade_result(
  pass_if(~identical(.result, df1[5,3] )),
  
  correct="תשובה נכונה", 
  incorrect = "נסו שוב"
)
```



שורות הקוד הבאות מבצעות השמה מחדש לערך של שורה 4, עמודה 3: ציון תחושת המסוגלות
העצמית של הנבדק בשורה הרביעית. 

הקוד משנה את ערך זה כך שיהיה שווה ל2 ומדפיס את הטבלה המעודכנת. שימו לב שהציון של הנבדק אכן השתנה ל2.


`תרגיל:` שנו את עיר מגוריו של הנבדק בשורה השנייה לDimona.


```{r exercise3_5, exercise=TRUE, exercise.eval = FALSE,  exercise.setup = "prepare-df1"}

df1[4,3] = 2
print(df1)
```

```{r exercise3_5-check}
library(dplyr)
grade_result(
  pass_if(~identical(.result, df1 %>%
                       mutate(city= ifelse(subject == 2,
                                           "Dimona",
                                           city)))),
  
    correct="מצוין", 
  incorrect = "נסו שוב"
)
```




#### חיתוך וקטורים ותת-טבלאות

ניתן להשתמש בסוגריים מרובעים גם כדי לחתוך תת-טבלה מתוך הטבלה המקורית.
בחיתוך כזה במקום לתת מספר בודד כאינדקס עבור השורות / העמודות נשתמש
בוקטור המציין את מספרי השורות / עמודות שאנחנו רוצים לחתוך.


כשמדובר בעמודות, אפשר לציין את העמודות שברצוננו לחתוך גם באמצעות שמות
העמודות. במקום לשים באינדקס העמודות מספר, נספק את שם העמודה או וקטור של
שמות העמודות אותן נרצה לחתוך.

צורה נוספת, ואחרונה לשיעור זה, בה אפשר לבצע חיתוך נעזרת בהשארת אחד האינדקסים ריק. במקרה כזה נקבל את כל השורות / עמודות של הטבלה.


לפניכם מספר דוגמאות:

```{r slicing1,  exercise=TRUE, exercise.eval = FALSE, exercise.setup = "prepare-df1"}

# שורות 2,4 ו 6 מעמודות 1 ו 2
df2 = df1[c(2,4,6),1:2] 
print(df2)

```

```{r slicing2,  exercise=TRUE, exercise.eval = FALSE, exercise.setup = "prepare-df1"}

# שורה 1, כל העמודות
row1 = df1[1,] 
print(row1)

```

```{r slicing3,  exercise=TRUE, exercise.eval = FALSE, exercise.setup = "prepare-df1"}

# כל השורות, עמודה 2 ו3
columns_2_and_3 = df1[,c(2,3)] 
print(columns_2_and_3)
```

```{r slicing4,  exercise=TRUE, exercise.eval = FALSE, exercise.setup = "prepare-df1"}

 # עיר המגורים של הנבדקים בשורות 3 ו1
city_subj_3_and_1 = df1[c(3,1),"city"]
print(city_subj_3_and_1)
```

```{r slicing5,  exercise=TRUE, exercise.eval = FALSE, exercise.setup = "prepare-df1"}
# כל השורות - אבל רק את העמודות של מספר ההרצה וציון תחושת המסוגלות העצמית
subj_and_self_efficacy = df1[,c("subject","self_efficacy_score")]  
print(subj_and_self_efficacy)

```

### תרגול

`תרגיל` חתכו מתוך הטבלה `df1` טבלה המכילה את מספר הנבדק וציון המסוגלות העצמית של נבדקים 1 ו5.


```{r exercise4_5_slicing_multi_index, exercise=TRUE, exercise.eval = FALSE,  exercise.setup = "prepare-df1"}
print(df1)

df2 =  # חיתוך מתוך df1

print(df2)
```

```{r exercise4_5_slicing_multi_index-check}
grade_result(
  pass_if(~identical(.result, df1[c(1,5),c(1,3)])),
  
    correct="יפה מאוד!", 
  incorrect = "נסו שוב"
)
```


### שליפת עמודות

ניתן לשלוף עמודות גם באמצעות פקד שליפת העמודות המיוצג על ידי התו `$`.

למשל: הקוד `df1$city` ישלוף את עמודת הערים מהטבלה `df1`.

כמו בחיתוכים הקודמים, גם כאן נוכל להשתמש בחיתוך גם כדי לערוך את הערכים הקיימים בטבלה:

במידה ונשתמש בפקד כשהוא מלווה בשם של עמודה הקיימת בטבלה - הוקטור שנכניס ידרוס את העמודה ויחליף אותה. במידה ונשתמש בשם של עמודה שלא נמצאת בטבלה - תיווצר עמודה חדשה בשם זה.


בדרך זו ניתן לייצר עמודות חדשות, בין היתר גם על סמך עמודות קיימות, או לבצע תיקונים
בעמודות הקיימות.


לפניכם מספר דוגמאות. הריצו את הקוד בשביל לראות מה קורה:

```{r df5, exercise=FALSE, exercise.eval = FALSE, , exercise.setup = "prepare-df1"}

# חיתוך עמודת הערים מתוך הטבלה
print(df1$city)

# יצירת עמודה חדשה המכילה נתוני גיל עבור הנבדקים
df1$age = c(21,35,42,32,24,19)
print(df1)

# דריסת עמודת הערים
df1$city = c("Haifa","Be'er-Sheva","Jerusalem","Be'er-Sheva","Jerusalem","Acre")
print(df1)

# יצירת עמודות על סמך עמודות קיימות

# העברה של ציוני המסוגלות העצמית לסולם 1-100 במקום 1-5
df1$se_scaled_1_to_100 = (df1$self_efficacy_score-1) * 25 
print(df1)

# חישוב מספר הנקודות בהן כל נבדק גבוה או נמוך מהממוצע בציון המסוגלות העצמית 
df1$se_relative_to_mean = df1$self_efficacy_score  -  mean(df1$self_efficacy_score)  
print(df1)

```

`תרגיל:` הוסיפו לטבלה עמודה בשם `high_self_efficacy` המכילה את הערך `TRUE` אם ציון תחושת המסוגלות העצמית של הנבדק גבוה מ2.5 ו`FALSE` אם הוא נמוך או שווה ל2.5.


```{r exercise5_5_dolar, exercise=TRUE, exercise.eval = FALSE,  exercise.setup = "prepare-df1"}

# הוספת עמודה
        
# הדפסת הטבלה לאחר הוספת העמודה
print(df1)                                 

```

```{r exercise5_5_dolar-check}
library(dplyr)

grade_result(
  pass_if(~identical(.result,
                     df1 %>% mutate(high_self_efficacy = self_efficacy_score > 2.5))),
  
    correct="נהדר!", 
  incorrect = "נסו שוב"
)
```



## טעינת טבלאות

עד עכשיו הגדרנו את הטבלאות שלנו בעצמנו. בפועל - רוב הטבלאות איתן נעבוד
בR יכילו נתונים שנאספו על ידי תוכנות אחרות. כדי לנתח אותן בR נצטרך קודם
כל לייבא את הנתונים לסביבת העבודה שלנו.

ישנן מספר פונקציות בR אשר מאפשרות לנו לטעון קבצים מפורמט csv או xlsx
(שני הפורמטים הנפוצים לשמירת נתונים המיוצגים בטבלה), אך הדרך הידודותית
ביותר למשתמש נעזרת בכלי של Rstudio שנועד להקל על ייבוא הקבצים.

כדי לגשת לכלי הייבוא, פתחו את Rstudio ובחרו בלשונית "file" שבסרגל הכלים
שבראש החלון של התוכנה. מכאן בחרו ב "Import Dataset" ובאחת מ2 האפשרויות
הראשונות ("From Text") במידה ואתם מעוניינים לטעון קובץ מפורמט csv, או בשלישית במידה והקובץ
שלכם הוא מפורמט xslx בו נשמרים קבצי אקסל.

![](images/import.jpg){width="200"}   

ביחידה זו נתרגל טעינת קובץ מפורמט csv ולכן נבחר באפשרות השנייה - "From
Text (readr)...". גם האופצייה הראשונה אפשרית, אבל השנייה נוחה יותר.
כשנבחר באפשרות זו ייפתח לנו חלון שיאפשר לנו לבחור קובץ לייבוא, לשלוט
בכמה אפשרויות בתהליך הייבוא ולצפות בתצוגה מקדימה של הטבלה שאנחנו
מייבאים.

תהליך הייבוא קורה בכמה שלבים:

::: text-block1
**1) בחירת הקובץ:**

נלחץ על כפתור ה"Browse..." שבפינה הימנית עליונה של החלון. ייפתח עבורנו
חלון בו נוכל לחפש ולבחור את הקובץ שאנו מעוניינים לייבא.

**2) התאמת אפשרויות הייבוא:**

החלון מאפשר לנו להתאים את הייבוא לאופן בו הנתונים שמורים בקובץ שלנו.
למשל, הוא מאפשר לנו לבחור האם השורה הראשונה בקובץ שאנחנו מייבאים מכילה
את הכותרות של העמודות ולדלג על מספר שורות במידה והקובץ מכיל כמה שורות
שאינן מכילות נתונים בתחילת הקובץ (באמצעות שדה ה"Skip"). בנוסף, החלון
מאפשר לנו גם לקבוע מה יהיה שם האובייקט בו נרצה לשמור את הטבלה שאנחנו
מייבאים, באמצעות השדה "Name" ולערוך את סוג הנתונים בכל עמודה.

**3) ייבוא הקובץ:**

את הייבוא עצמו אפשר לבצע ב2 דרכים שונות:

הדרך הראשונה היא באמצעות לחיצה על כפתור "Import" שבתחתית החלון. כפתור זה
יבצע את שורות קוד המתאימות להעדפות שהגדרנו בחלון. ניתן לראות את שורות
קוד אלו בחלונית בתחתית החלון.

לרוב, לא נרצה רק לייבא את הטבלה לסביבת העבודה הנוכחית שלנו, אלא נרצה גם
להעתיק את שורות הקוד שמייבאות את הקובץ לקובץ הסקריפט שלנו, כך שבהרצה
עתידית של הקובץ הקוד ייבצע את ייבוא הטבלה באופן אוטומטי. כדי לעשות זאת
נעתיק את שורות הקוד מחלונית הקוד ונדביק אותן בקובץ הסקריפט שלנו - ואז
נריץ אותן מתוך הקובץ.

שורת הקוד הראשונה מייבאת את החבילה שנדרשת לצורך הייבוא, השנייה מבצעת את
הייבוא בפועל והשלישית פותחת עבורנו את הטבלה שייבאנו לצפייה. כשאנו
מעתיקים את הקוד לסקריפט בדרך כלל נעדיף להשמיט את השורה השלישית.
:::




![](images/Rstudio logo.png){width="100"}  

`תרגיל:` ייבאו את הקובץ "need_for_cognition.csv" שבאתר המודל של הקורס.
לסביבת העבודה שלכם בRstudio
[(קישור לקובץ במודל)](https://moodle.huji.ac.il/2024-25/mod/resource/view.php?id=135118)

ראשית, הוריד את הקובץ ושמרו אותו בתיקייה לבחירתכם, ואח"כ היעזרו בחלון
הייבוא כדי לייצר את שורות הקוד שיאפשרו ייבוא תקין של הטבלה לסביבת העבודה
שלכם. הריצו את שורות הקוד וודאו שהייבוא התנהל באופן תקין.

קובץ נתונים זה מתאר את ציוניהם של 2466 משתתפים בשאלון המודד "need for
cognition" - מידת הצורך שאנשים מרגישים להבין את העולם סביבם. בהמשך השיעור
נציץ אל תוך טבלה זו ונבין את המבנה שלה.

```{r question2_5, echo = FALSE}
question("מה הגיל של הנבדקת הראשונה בקובץ שייבאתם?",
         type = "learnr_text",
         answer("19", correct = TRUE),
         allow_retry = TRUE,
           
  correct="מעולה", 
  incorrect = "נסו שוב"
)
```

### ייצוא טבלאות

לפני שנמשיך להצצה אל תוך התוכן של הטבלה, ניגע לרגע בייצוא טבלאות.

ישנם מספר דרכים לייצא נתונים מתוך R, כדי לשתף אותם עם אנשים אחרים או לשמור אותם לטובת ניתוחים עתידיים.  דרך אחת שכזו, המתאימה לשיתוך טבלאות, ממירה את הטבלה שייצרנו בR ושומרת אותו כקובץ csv.

למשל, ייתכן שייבאנו טבלה, ערכנו אותה - ואנחנו רוצים לשתף את הטבלה הערוכה
עם אנשים אחרים. כדי לשתף אותה נשמור אותה חזרה לקובץ csv אותו נוכל לשלוח
למי שנרצה.

:הייצוא נעשה באמצעות הפקודה הבא:

```{r export, eval = FALSE}

write.csv(df1, "export dataframe example.csv")

```

בדוגמה זו ייצאנו את הטבלה df1 שנמצאת בסביבת העבודה שלנו ושמרנו אותה בתור
קובץ CSV בשם "export dataframe example.csv".

הקובץ יישמר בתיקיית העבודה שלכם (working directory- תיקיית ברירת המחדל שאליה תיגש התוכנה כדי לייבוא ולשמור קבצים).


ניתן לברר מה היא תיקיית העבודה הנוכחית ולשנות אותה באמצעות הפונקציות הבאות:


```{r getwd, eval = FALSE}

# מציאת כתובת העבודה
working_directory = getwd()
print(working_directory)

# שינוי כתובת העבודה
# כרגע אנחנו מגדירים את התיקייה לפי התיקייה הנוכחית, ולכן היא לא תשתנה. החליפו את  
# "C:\Users\one\Desktop\R files\file directory example"
# בכתובת לתיקייה אחרת כדי לשנות את תיקיית העבודה, במידה ואתם מעוניינים בכך

new_working_directory = "C:\Users\one\Desktop\R files\file directory example"
setwd(new_working_directory)

```

![](images/Rstudio logo.png){width="100"}  

`תרגיל:`
ייבאו את הטבלה עם תשובות הנבדקים בשאלון "need for cognition"אל סביבת העבודה שלכם בRstudio, חתכו ממנו תת טבלה המכילה את התצפיות ב10 השורות הראשונות עבור העמודות "session_id" ו"age" וייצאו את הטבלה החדשה כקובץ csv.

פתחו את הקובץ שנוצר וודאו שהוא מכיל את הנתונים הנכונים.

```{r question3_5, echo = FALSE}

question(
  "כמה עמודות היו בקובץ שפתחתם?",
  answer("אחת",  message = "תשובה שגויה. נסו שוב"),
  answer("שתיים", 
         message = "הטבלה שיצרתם אמנם הכילה שתי עמודות, אבל בדקו כמה יש בקובץ ששמרתם." ),
  answer("שלוש",
         correct = TRUE,
         message = "תשובה נכונה. תהליך השמירה הוסיף עמודה ללא כותרת המכילה את המספר הסידורי של כל שורה."),
  answer("ארבע",  message = "תשובה שגויה. נסו שוב"),
  allow_retry = TRUE
)

```


## בחינה ראשונית של טבלה


**ייבאנו טבלת נתונים - מה עכשיו?**


בדרך כלל אחרי שנייבא טבלת נתונים נצטרך להזכר במבנה שלה.

בין אם הטבלה מכילה נתונים שאספנו בעצמנו או נתונים שקיבלנו מגורם אחר, לפני שנוכל להתחיל לעבוד עם הנתונים נצטרך לבדוק איך קוראים לעמודות, מה כל עמודה מייצגת, באיזה אופן הנתונים מקודדים, כמה תצפיות הטבלה מכילה וכו'

יש כמה פונקציות בR שיכולות לעזור לנו בכך.



### מימדי הטבלה

המאפיינים הכי בסייסים שיעניינו אותנו נוגעים לגודל הטבלה - כמה שורות ועמודות היא מכילה.

שורות הקוד הבאות מבררות את מימדי הטבלה. 


```{r}
# מספר השורות והעמודות שבטבלה שייבאנו
dim(need_for_cognition)  

# מספר השורות
nrow(need_for_cognition)

# מספר העמודות
ncol(need_for_cognition)
```

התוצאות יספקו לנו מידע כללי לגבי כמות התצפיות ומספר המדדים השונים שהטבלה מכילה.

```{r question4_5, echo = FALSE}

question_text(
  "כמה תצפיות מכילה הטבלה שייבאתם?",
  answer("2466", correct = TRUE),
  answer("10", message = ".התשובה אינה נכונה. הזנתם את מספר העמודות בטבלה, במקום את מספר השורות"),
  allow_retry = TRUE,
  
    
  correct="תשובה נכונה", 
  incorrect = "נסו שוב"
)

```

פונקציות אלו שימושיות במיוחד כשאנחנו צריכים להתייחס למימדים אלו בקוד, כך שבמקום לבדוק בעצמנו את אורך הטבלה ולהקליד את המספר המתאים נוכל להשתמש בקוד שיעשה את הבדיקה הזו עבורנו.

למשל, בקטע הבא אנחנו נעזרים בפונקציה `nrow` כדי לייצר וקטור עם המספרים בין 1 למספר השורות שקיים בטבלה שלנו ומשתמשים בו כדי להוסיף לטבלה עמודה עם מספר סידורי עבור כל תצפית.

```{r}
need_for_cognition$subject_index = 1:nrow(need_for_cognition)
```


### צפייה בתמצית הטבלה, שמות העמודות ובשורות הראשונות / אחרונות


בנוסף, ישנן מספר פונקציות שיאפשרו לנו לבצע הצצה ראשונית אל תוך **התוכן** של
הטבלה והעמודות שהיא מכילה. הן מאפשרות לנו להבין באופן גס מה כלול בכל עמודה ולהיזכר בכותרות של העמודות השונות.


האופציה הנגישה ביותר כוללת צפייה בטבלה באמצעות לחיצה על שמה בחלונית הסביבה (Environment) המופיעה בצד הימני של המסך בRstudio. כאשר נלחצת על שם הטבלה תוצג תצוגה מקדימה של הטבלה והעמודות שבה.

במידה ונרצה לצפות רק בתחילתה או סופה של הטבלה נוכל להשתמש בפונקציות הבאות:

```{r}

# שליפת השורות הראשונות בטבלה 
head(need_for_cognition)# ברירת המחדל היא 6 שורות, אך ניתן גם לבקש מספר אחר של שורות
head(need_for_cognition,4)

# שליפת השורות האחרונות בטבלה
tail(need_for_cognition) 

```

בנוסף, ייתכן ונרצה לקבל תקציר על כל עמודה בטבלה, כך שנוכל להבין מה כל עמודה מייצגת ואילו ערכים היא מכילה. כדי לעשות זאת נשתמש בפונקציה `summary` שתחזיר לנו סיכום לכל עמודה בהתאם לסוג הנתונים המיוצגים בה.

``` {r}
# סיכום בסיסי של כל עמודה בטבלה, בהתאם לסוג העמודה
summary(need_for_cognition)

```

במידה ונרצה להיזכר רק בשמות העמודות או בצורה הספציפית בה הן מאויתות נוכל להשתמש בפונקציה `names`.

```{r}
# שליפת שמות העמודות 
col_names = names(need_for_cognition)
print(col_names)

```



```{r question5_5, echo = FALSE}

question(
  "<div style='direction: rtl;'> בקובץ הנתונים שטענתם - מה מכילה העמודה attention_check?</div>",
  answer("<div style='direction: rtl;'>את ציון הנבדק במטלת קשב (ציון רציף 1 עד 100)</div>" , message = "בדקו את הערכים שהעמודה מכילה - מה הם יכולים לייצג?"),
  answer("<div style='direction: rtl;'>את סוג מבחן הקשב שהנבחן עבר (רגיל / קל / קשה)</div>" , message = "בדקו את הערכים שהעמודה מכילה - מה הם יכולים לייצג?"),
  answer("<div style='direction: rtl;'>האם הנבחן עבר את מבחני הקשב (1 = עבר, 0 = לא עבר)</div>", correct = TRUE , message = "נכון מאוד"),
  answer("<div style='direction: rtl;'>האם הנבחן עבר את מבחני הקשב TRUE = עבר, FALSE = לא עבר)</div>" , message = "בדקו את הערכים שהעמודה מכילה - מה הם יכולים לייצג?"),
  allow_retry = TRUE
)

```

## תרגיל

![](images/Rstudio logo.png){width="100"}   

העתיקו את הקוד הבא לסביבת העבודה שלכם והשלימו אותו בהתאם להערות בגוף הקוד.


בסיום העבודה, העתיקו את הקוד שכתבתם חזרה לחלונית הקוד, למעט שורות הקוד העוסקות בייבוא הנתונים (הטבלה כבר טענה מראש, תחת השם "need_for_cognition").


```{r summary_exercise, exercise=TRUE, exercise.eval = FALSE, , exercise.setup = "prepare-df1"}
library(readr)

# הורידו את הקובץ מהמודל וייבוא אותו לסביבת העבודה שלכם
need_for_cognition = ___________________    

# ~~~ מחקו את כל השורות  מעל נקודה זו בהעתקה חזרה לחלונית הקוד   ~~~ 

# הדפיסו את שמות העמודות בטבלה
column_names = ___________________________
print(column_names)


# כעת, צרו טבלה נוספת המכילה את 100 השורות הראשונות של 
# מספרי הנבדקים והתשובות שלהם לשאלון (עמודות 1 ו5 עד 10)
table1 = _____________________________

# וודאו שלטבלה החדשה יש את מספר השורות והעמודות הנדרשות (100 שורות 7 עמודות)
number_of_columns = _____________
number_of_rows = ______________

print(paste("table1 has",number_of_columns,"columns and",number_of_rows,"rows"))

# בטבלה ישנן 2 עמודות המציגות ציונים עבור שאלות שהינן בכיוון ההפוך ביחס לשאר השאלון 
# לפני שנוכל לחשב את המדד המשוכלל עבור השאלון נצטרך להפוך אותן חזרה

# צרו 3 עמודות חדשות עבור שאלות 02, 03 ו 06 שיכילו את הציון ההפוך לציון בעמודות המקוריות
# מאחר וציוני השאלון נעים בין 1 ל5, על מנת להפוך את הציון עלינו להחסיר את הציון המקורי מ6, כלומר
# new column = 6 - old column

# מומלץ להיעזר בפקד ה$ כדי להגדיר את העמודות החדשות ולשלוף את ערכיהן של העמודות הישנות
# קראו לעמודות החדשות לפי שמות העמודות הישנות ללא האות r 
  
table1$nfc_02 = ____________________
table1$nfc_03 = ____________________
table1$nfc_06 = ____________________


#  כעת, צרו עמודה חדשה שתכיל את סכום ששת העמודות עבור כל נבדק
table1$nfc = ___________________________________________________________________

# שלפו את עמודה זו ושמרו אותה באובייקט הבא:
nfc_score_column =  ____________

# 

```


```{r summary_exercise-solution}
# הדפיסו את שמות העמודות בטבלה
column_names = names(need_for_cognition)
print(column_names)

# כעת, צרו טבלה נוספת המכילה את 100 השורות הראשונות של 
# מספרי הנבדקים והתשובות שלהם לשאלון (עמודות 1 ו5 עד 10)
table1 = need_for_cognition[1:100,c(1,5:10)]

# וודאו שלטבל החדשה יש את מספר השורות והעמודות הנדרשות (100 שורות ו7 עמודות)
number_of_columns =  ncol(table1)
number_of_rows = nrow(table1)

print(paste("table1 has",number_of_columns,"columns and",number_of_rows,"rows"))

# בטבלה ישנן 2 עמודות המציגות ציונים עבור שאלות שהינן בכיוון ההפוך ביחס לשאר השאלון 
# לפני שנוכל לחשב את המדד המשוכלל עבור השאלון נצטרך להפוך אותן חזרה

# צרו 3 עמודות חדשות עבור שאלות 02, 03 ו 06 שיכילו את הציון ההפוך לציון בעמודות המקוריות
# מאחר וציוני השאלון נעים בין 1 ל5, על מנת להפוך את הציון עלינו להחסיר את הציון המקורי מ6
# new column = 6 - old column

# מומלץ להיעזר בפקד ה$ כדי להגדיר את העמודות החדשות ולשלוף את ערכיהן של העמודות הישנות
# קראו לעמודות החדשות לפי שמות העמודות הישנות ללא האות r 
  
table1$nfc_02 = 6 - table1$nfc_02r
table1$nfc_03 = 6 - table1$nfc_03r
table1$nfc_06 = 6 - table1$nfc_06r


# כעת, צרו עמודה חדשה שתכיל את סכום ששת העמודות עבור כל נבדק
table1$nfc = table1$nfc_01 +table1$nfc_02 +table1$nfc_03 +table1$nfc_04 +table1$nfc_05 +table1$nfc_06 

# שלפו את עמודה זו ושמרו אותה באובייקט הבא:
nfc_score_column = table1$nfc
```

```{r summary_exercise-check}
grade_this({
  
   if (user_object_get("number_of_columns") != 7){
     fail("מספר העמודות אינו נכון")}
   else if(user_object_get("number_of_rows") != 100) {
     fail("מספר השורות אינו נכון")}
   else if(!identical(user_object_get("nfc_score_column"), solution_object_get("nfc_score_column"))) {
     fail("טעות בחישוב עמודת הציונים או בשליפתה")}
   else {pass("עבודה מצוינת!")}
}
   )
```


## הגשה
  עברו על הקובץ וודאו שהגשתם את כל התרגילים ועניתם על כל השאלות
  
  במידה וכל התשובות שלכם תקינות יש ללחוץ על הכפתור:  Generate, להעתיק את הטקסט שמופיע בחלון למטה ולהגישו במודל  
  בהצלחה!

```{r context="server"}
learnrhash::encoder_logic()
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui()
```