---
title: "Lesson 8: Coding resources" 
output:
  learnr::tutorial:
    progressive: true
    allow_skip: false
    
description: >
 <div style='direction: rtl;'>
 השיעור נוגע בעזרים חיצוניים שניתן ומומלץ להשתמש בהם בעת ניתוח נתונים בR. נדבר על תיעוד נלווה, בלוגים ופורומים אינטרנטיים ושימוש בבינה מלאכותית.
 </div>
     
runtime: shiny_prerendered  
editor_options: 
  markdown: 
    wrap: 72
---

<style>
h1, h2, h3, h4, h5, h6 {
  direction: rtl;
}
p {
  direction: rtl;
}

.text-block1 {
  direction: rtl;       /* Set text direction to right-to-left */
  text-align: right;
  background-color: #f0f0f0; /* Light grey background */
  padding: 10px;
  border-radius: 5px;
  border: 1px solid #ddd; /* Light border */
  margin: 10px 0; /* Space around the block */
  
</style>

<!-- tutorial options and cheching options - TODO set params -->

```{r setup, include=FALSE}
library(learnr)
library(dplyr)
library(gradethis)
library(tidyverse)

tutorial_options(
  exercise.timelimit = 60
)


knitr::opts_chunk$set(error = TRUE)
```

```{r prepare-df1}
set.seed(111)
df1 = data.frame(subject = rep(1:5,each = 4),
                 trial   = rep(1:4,5),
                 score   = round(rep(1:4,5) + rnorm(20, 20,5)) %>% as.character())

```

## מבוא

במהלך הסמסטר עברנו על מגוון פונקציות וצורות ניתוח שנמצאות בשימוש תדיר
בניתוח נתונים והמחשה חזותית. פונקציות אלו מהוות את הבסיס לרוב הניתוחים
הנפוצים, אבל לאו דווקא יכסו את כל צרכי הניתוח שלכם.

בשיעור הנוכחי אנחנו נתייחס למקורות מידע חיצוניים, העשויים לעזור לכם
בניתוחים עתידיים.

אנחנו נתייחס לשלוש סוגי מקורות מומלצים:

**1)** תיעוד נלווה של פונקציות\
**2)** בלוגים ופורומים אינטרנטיים\
**3)** בינה מלאכותי

## תיעוד נלווה

המשאב הראשון והוותיק ביותר שעומד לרשותנו הינו התיעוד הנלווה של פונקציות
(Documentation). לכל חבילה בR מוצמד תיעוד מתאים - המפרט עבור כל פונקציה
את מטרתה, את הנתונים השונים שהיא יכולה או צריכה לקבל כדי לפעול באופן
תקין ומספר דוגמאות שימוש.

ניתן לגשת לתיעוד זה בשני דרכים:

### באמצעות הפקד `?`

כשנריץ סימן שאלה מלווה בשמה של פונקציה - יפתח עבורנו חלון העזרה "help"
המובנה בתוך Rstudio ובתוכה יופיע התיעוד הרלוונטי לפונקציה.

למשל - במידה ונרצה לברר איך משתמשים בפונקציה `round` נריץ את השורה הבאה:

```{r , eval = FALSE}
?round
```

הריצו את שורת הקוד בRstudio ועיינו בתיעוד הנלווה.

תוכלו לראות שהתיעוד מציג את אופן השימוש בפונקציה: היא מקבלת כאיבר ראשון
את X - שהוא האובייקט עליו נרצה להפעיל את הפונקציה, וכאיבר שני את מספר
הספרות אחרי הנקוב העשרונית אליהם נרצה לעגל (digits). במקרה זה התיעוד
כולל גם פירוט על מספר פונקציות דומות.

### Rdocumentation

התיעוד שמוצג בחלון העזרה מופיע גם באתר האינטרנט
[Rdocumentation](https://www.rdocumentation.org/ "קישור לאתר").

המידע באתר אינו שונה מחלון העזרה, אבל ההצגה שלו על פני דף שלם בדרך כלל
נוחה יותר מאשר התצוגה בחלונית.

בשני המקרים - המידע שמסופק הוא מאוד טכני ולא תמיד קל להבנה. הוא שימושי
בעיקר כדי להיזכר בנתונים שפונקציה אמורה לקבל וכדי לראות דוגמאות שימוש
בסיסיות.

`תרגיל:` חפשו ברשת או בחלונית העזרה את התיעוד הנלווה של הפונקציה
`substr`, המשמשת לשליפה של חלק מתוך מילה / משתנה תווי (character).

היעזרו בפונקציה זו על מנת לשלוף את שלושת התווים הראשונים בקוד המזהה של
הנבדקים ושמרו אותם בעמודה חדשה בשם "short_ID":

```{r rdocumentation, exercise=TRUE, exercise.eval = FALSE, exercise.setup = "prepare-df1"}

df = data.frame(ID = c("001p","002f","003f","004f","005p","006f"),
                score = c(45,65,57,71,52,43))

df = df %>% mutate(short_ID = substr(___________))

print(df$short_ID)

```

```{r rdocumentation-check}

df = data.frame(ID = c("001p","002f","003f","004f","005p","006f"),
                score = c(45,65,57,71,52,43))
  
  
grade_result(
  pass_if(~identical(.result,substr(df$ID,1,3) ))
)
```

## בלוגים אינטרנטיים

משאב נוסף, שהינו מעט יותר ידידותי למשתמש, כולל מידע שהועלה על ידי
משתמשים לרשת.

בחלק מהמקרים מדובר בפורומים בהם אנשים העלו בעיות שהם נתקלו בהם וקיבלו
תשובות שעשויות לעזור לנו בבעיות איתן אנחנו מתמודדים ובחלקם מדובר בבלוגים
מסודר, בהם אנשים מעלים מדריכים מפורטים לגבי אופן השימוש בפונקציות
מוסימות ודרכים בהם אפשר לבצע ניתוחים מורכבים. הפירוט כאן לרוב יהיה עשיר
יותר, לפעמים בליווי דוגמאות או סרטונים רלוונטיים.

את רוב הרשומות והשאלות בפורומים אפשר למצוא באמצעות חיפוש ברשת של הבעיה
איתה אתם מתמודדים, כך שאין חשיבות אמיתית להיכרות עם הבלוגים השונים מראש
אבל נמליץ כאן בכל זאת על בלוג אחד שעשוי להיות שימושי:

::: text-block1
[The R Graph Gallery](https://r-graph-gallery.com/ "קישור לאתר")- בלוג
המפרט על אופן היצירה של סוגים שונים של תרשימים, רובם באמצעות `ggplot`
ומספר חבילות המאפשרות לייצר תרשימים אינטראקטיביים.
:::

למיטיבי לכת ישנם גם פורומים שבועיים המעדכנים בדבר שינויים ועדכונים
בתוכנה ואפילו [פודקסט](https://r-podcast.org/).

## בינה מלאכותית

### מבוא

עד לאחרונה, השימוש ברשומות ופורומים היווה את המשאב העיקרי עבור למידה
ופתרון בעיות בR, אבל מאז העלייה של מנועי הבינה המלאכותית בשנים האחרונות
רבים עוברים לעשות שימוש במנועים כדוגמת chatGPT, Claude וכו'. היתרון
בשימוש בבינה מלאכותית הוא בכך שהם מאפשרים לשאול שאלות תכנותיות באופן
ישיר (חלקם גם תומכים בשאלות בעברית) ולקבל תשובה מפורטת המותאמת היטב למה
שאנחנו זקוקים לו. התוכנות מסוגלות גם לייצר קטעי קוד, קצרים או ארוכים,
כדי להדגים את התשובות שלהם ולתקן בעיות בקוד שלנו.

הכלים האלו אמנם מגבירים את היכולת שלנו לפתור בעיות ולבצע ניתוחים
מורכבים - אבל חשוב להכיר גם את המגבלות שלהם.

![](images/ai.jpg){width="300"}

בפרק הנוכחי נדבר על מספר שימושים מומלצים בבינה מלאכותית ונתעכב על כמה
אזהרות חשובות.

### בירור לגבי פונקציות

השימוש הנפוץ ביותר בבינה מלאכותית, בו הוא גם בדרך כלל מצטיין, כולל
התייחסות אליו כאל מאגר מידע. במקום לנסות מילות מפתח שונות ולגלול ברחבי
הרשת עד שתמצאו את הפונקציה הספציפית שאתם צריכים או רשומה נוחה לגבי אופן
השימוש בה - התוכנה בדרך כלל תדע לספק לכם תשובות מהירות ומבוססות לגבי
איזו פונקציה מבצעת את הפעולה שאתם צריכים או לגבי הדרכים השונות להשתמש
בפונקציה נתונה.

::: text-block1
גם כאן יכולות להיות טעויות והתכנה לעיתים רחוקות תמציא פונקציות או תשמשתש
בהן בצורה לא נכונה, אבל ברוב המקרים הפשוטים היא מספקת מענה מוצלח מאוד.
בכל אופן - נוכל לבדוק את התשובות שלה באופן מיידי, או באמצעות הרצת
הפונקציה או באמצעות עיון בתיעוד של הפונצקיה וברשומות קיימות לגביה.
:::

`תרגיל` : בררו איזו פונקציה יכולה לשמש אותנו כדי להפריד עמודה לשניים.

בטבלה הבאה ישנה עמודה בשם "session". הפרידו אותה לעמודות
"subject"ולעמודה "trial" כך שכל עמודה תכיל את התוכן המתאים.

למשל - בשורה הראשונה נרצה שתחת העמודה "subject" יופיע "s01" ותחת העמודה
"trial" יופיע "t1".

```{r separate, exercise = TRUE}

df = data.frame(session = c("s01_t1","s01_t2","s01_t3","s02_t1","s02_t2" ))

df = df %>% ________________
  
print(df)

```

```{r separate-check}

grade_result(
  pass_if(~identical(.result, 
                     
data.frame(session = c("s01_t1","s01_t2","s01_t3","s02_t1","s02_t2" ))%>%
  separate(session, into = c("subject","trial"))
  
  ))
)
```

### פתרון תקלות

אחד העיסוקים הכי מתסכלים בתכנות בכלל, ובניתוח נתונים בR בפרט, הוא שמדי
פעם אנחנו נתקלים בתקלות ובעיות בהרצה שלא תמיד ברור לנו איך עלינו לפתור
אותן. הפלט בR יכול להכווין אותנו לפונקציה שבה ארעה התקלה - אבל גם אז
לפעמים קשה להבין ממה התקלה נובעת.

מנועי בינה מלאכותית יכולים לעזור לנו לצאת מהבוץ במקרים רבים שכאלו. אם
נדביק את הפלט של השגיאה למנוע הבינה המלאכותית הוא יוכל להסביר לנו בפירוט
את המשמעות של השגיאה שחווינו. הוא לא תמיד יוכל להסיק את שורש הבעיה, אבל
הוא יוכל לתת לנו הבנה ראשונית של הבעיה לכל הפחות.

אם נצרף לפנייה שלנו גם את החלק הבעייתי בקוד שלנו, התוכנה תוכל להבין
בצורה טובה יותר את מקור הבעייה ובמקרים רבים גם להציע לנו איך לתקן אותה
או לספק לנו קטע קוד מתוקן.

`תרגיל` הריצו את קטע הקוד הבאה והיעזרו במנוע בינה מלאכותית לבחירתכם כדי
להבין את מקור השגיאות ולתקן אותן. שימו לב - הקוד מכיל מספר שגיאות שונות.

```{r warnings, exercise = TRUE, exercise.setup = "prepare-df1"}
# df1:
# הטבלה מתארת את הציונים של נבדק בכמה חזרות שונות של אותה המטלה

# subject - המספר המזהה של הנבדק
# score  - ציון הנבדק
# trial   - מספר החזרה

df1 %>% head() %>% print() # הדפסת השורות הראשונות

df1_summary = df1 %>% 
  group_by(subject) %>%                      
  summarise(trial_amount = max(trial), # מספר החזרות שביצע
            score = mean(score),       # הציון הממוצע 
            cumulative_score = sum(score))    # הציון המצטבר

cumulative_score = df1_summery %>% pull("cumulative_score")

print(cumulative_score)

```

```{r warnings-check}
set.seed(111)


grade_result(
  pass_if(~identical(.result,
                     
                     
                     
 data.frame(subject = rep(1:5,each = 4),
            trial   = rep(1:4,5),
            score   = round(rep(1:4,5) + rnorm(20, 20,5)))  %>% 
  group_by(subject) %>%                       
  summarise(trial_amount = max(trial, na.rm = T), # מספר החזרות שביצע
            cumulative_score = sum(score),
            score = mean(score, na.rm = T)     # הציון הממוצע 
            )  %>% pull(cumulative_score) ))
)
```

### יצירת קטעי קוד

שימוש נוסף, יעיל אך מעט בעייתי, בבינה מלאכותית הוא יצירה של קטעי קוד
שלמים. במקום להקדיש זמן יקר לכתיבת קוד, זיהוי שגיאות ותיקונים, אנחנו
לכאורה יכולים לבקש מהבינה המאכותית לכתוב את הקוד עבורנו.


יש בכך 3 בעיות עיקריות:

::: text-block1

**1) גם תוכנות טועות**

טעויות קטנות בקוד יגרמו לקוד כולו לא לעבוד - ובמידה והקוד לא נכתב על ידינו - יהיה
לנו מאוד קשה להבין את מקור הטעות ולתקן אותה. נוכל להיעזר בבינה המלאוכתית
בשביל תיקונים כאלו, אבל במקרים רבים נמצא את עצמנו עם שמיכת טלאים של קטעי
קוד שאנחנו פשוט לא מצליחים לתקן. זה קורה בעיקר בקטעי קוד ארוכים (מעל
פסקה אחת). נכון לרגע כתיבת מילים אלו - יש לא מעט תחומים ופעולות שמנועי בינה מלאכותית מסתבכים בהם. במקרים כאלו, גם אם תבקשו מהם לתקן את הקוד שקיבלתם מהם - הם לא יידעו לתקן את הבעיות.


כפועל יוצא מכך - מומלץ להגביל את השימוש בקוד שנוצר על ידי בינה מלאכותית לפסקה או פעולה אחת בכל פעם. אפשר ורצוי לבקש מהתוכנה שתסביר מה כל
חלק עושה - אבל כל עוד לא הבנו מה כל שורה עושה, השימוש וההתבססות על קטעי
קוד כאלו הוא הימור בעייתי. זה אמנם לפעמיםחוסך זמן בטווח הקצר, אבל במקרים רבים יוסיף
לכם הרבה עבודה בהמשך.

**2) קצרים בתקשורת**

במובנים רבים, האפשרות הזו יותר מסוכנת לנו מאשר יצירת קוד שלא עובד. במידה
והתכנה לא תבין את הניתוח / העריכה שאתם מתכננים היא יכולה לספק לכם קוד
מפורט ותקין שמבצע ניתוח כלשהו, שאינו הניתוח שאתם רוצים. כלומר - בנוסף
ל"השגחה" על תקינות הקוד מבחינה תכנותית, חשוב גם לוודא שהקוד באמת מבצע את
הניתוח שאנחנו מעוניינים בו. זו סיבה נוספת להתעקש להבין כל שורה בקוד שקיבלתם מבינה מלאכותית.

**3) יצירת תלות**

בשלב זה אתם עוד לומדים להשתמש בתוכנה. כלומר - יש הרבה פונקציות ופרקטיקות
ניתוח בסיסיות שאתם עוד רוכשים. מצד אחד - אלו תחומים בהם הבינה המלאכותית
בדרך כלל שולטת היטב ותוכל לספק תשובות טובות, אבל מצד שני - אם לא תלמדו
אותם, תתרגלו ותתנסו בהם, ייקח הרבה זמן עד שתוכלו לבצע את ניתוחים אלו
בעצמכם. התלות בבינה מלאכותית בנושאים אלו תעכב ותסרבל את העבודה שלכם.

בשלב זה, וגם בהמשך העבודה עם בינה מלאכותית, מומלץ להשתמש בה רק כמורה -
שמסביר, מדגים ועוזר לכם למצוא את הטעויות שלכם ולתקן אותן - אבל לא מבצע
את הניתוחים במקומכם.

:::


```{r}


```


### קופיילוט ![](images/copilot.png){width="50" height="26"}

תצורה נוספת ושימושית במיוחד של בינה מלאכותית הינה Github Copilot. היישום
כולל תוסף לRstudio הזמין החל מהגרסאות האחרונות של התוכנה, ומציע השלמות
לקוד בזמן הכתיבה.

לתוסף יש גישה לקובץ שאתם עובדים עליו ולקבצים אחרים שכתבתם, ובאמצעותם הוא
מנסה לנבא את מטרת הקוד שאתם כותבים ולהשלים אותו בהתאם.

מאחר והתוסף לא דורש מאיתנו להעתיק את הקוד לממשק אחר ולנסח באופן מפורש את
השאלות שלנו הוא הרבה יותר מהיר ונח לשימוש.

נכון להיום, התוסף דורש מנוי בתשלום, אך הוא מוצע חינם לסטודנטים דרך תכנית
Github education.

## הגשה

עברו על הקובץ וודאו שהגשתם את כל התרגילים ועניתם על כל השאלות

במידה וכל התשובות שלכם תקינות יש ללחוץ על הכפתור: Generate, להעתיק את
הטקסט שמופיע בחלון למטה ולהגישו במודל\
בהצלחה!

```{r context="server"}
learnrhash::encoder_logic()
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui()
```
